---
title: RStudio
description: 'Allgemeine Fragen zu RStudio'
---

## Anwendungsbereich R-Studio

```yaml
type: PureMultipleChoiceExercise
key: 85cda17d46
xp: 50
```

In welchen der folgend genannten Fälle ist RStudio eine ideale Anwendung?

`@hint`
Vielleicht solltest du ganz kurz mal die Homepage von R-Studio ansehen

`@possible_answers`
- [Um R-Programme zu schreiben, den Programmablauf zu testen und projektbasiert zu arbeiten.]
- Um die Ausgabe von R-Programmen tabellarisch und graphisch darzustellen.
- Um Graphiken und Tabellen von R-Programmen APA-konform darzustellen.
- Um R-Programme in lauffähige .exe-Dateien umzuwandeln.

`@feedback`
- Nicht schlecht für den Anfang!
- Nein, dafür kann es verwendet werden, aber es ist nicht die ideale Anwendung.
- Kann man machen, aber dafür ist R-Studio nicht gedacht.
- R ist eine Interpreter-Sprache, die Programme müssen daher nicht kompiliert werden.

---

## IDE

```yaml
type: PureMultipleChoiceExercise
key: c3ae5215ff
xp: 50
```

**Für was steht IDE im Zusammenhang mit RStudio?**

RStudio ist eine IDE Anwendung, mit der R einfacher verwendet werden kann. Es integriert dabei einige hilfreiche Entwicklungsumgebungen in eine Anwendung.

`@hint`


`@possible_answers`
- Intensive Design Environment
- Integrated Document Environment
- Independent Developer Ecosystem
- [Integrated Development Environment]

`@feedback`
- Nope
- Fast
- Nein
- Ja, das ist es!

---

## Kompatible Programme

```yaml
type: PureMultipleChoiceExercise
key: 225ae2d2bd
xp: 50
```

Welche Programme lassen sich nicht nur als Alternative zu R verwenden, sondern lassen sich auch in R einbinden?

`@hint`
- Es gibt zwei richtige Antworten.
- Ein Programm lässt sich mit dem Paket _jmv_ in R einbinden.

`@possible_answers`
- [JAMOVI]
- JPEG
- [JASP]

`@feedback`
- Perfekt!
- Leider nein, bei JPEG geht's um Bilder. 
- Richtig!

---

## Groß- und Kleinschreibung

```yaml
type: PureMultipleChoiceExercise
key: 62e5def0d9
xp: 50
```

In R wird zwischen Groß- und Kleinschreibung unterschieden, richtig oder falsch?

`@hint`


`@possible_answers`
- [richtig]
- falsch

`@feedback`
- Genau, nur weiter so!
- Nein, leider ist die Antwort _falsch_ wirklich falsch

---

## Dezimaltrennzeichen

```yaml
type: PureMultipleChoiceExercise
key: f9ca438ffc
xp: 50
```

Welches Dezimaltrennzeichen wird in R verwendet?

`@hint`


`@possible_answers`
- Komma
- [Punkt]
- Semikolon
- Colon

`@feedback`
- Nope
- Super, das ist es!
- Nope
- Nope

---

## Objektbezeichnung

```yaml
type: PureMultipleChoiceExercise
key: b43881f74b
xp: 50
```

Bei der Bezeichnung von Objekten gelten bestimmte Regeln. Welche der angeführten Regeln ist richtig?

`@hint`


`@possible_answers`
- Leerzeichen in Objektnamen sind erlaubt, aber mathematische Operatoren dürfen nicht verwendet werden.
- Alles ist erlaubt, außer dass Objektnamen mit einer Ziffer beginnen.
- [Keine Leerzeichen in Objektnamen, keine mathematischen Operatoren, erstes Zeichen muss ein Buchstabe sein.]

`@feedback`
- leider falsch
- nein, nicht richtig
- Yessss

---

## Ctrl-L

```yaml
type: PureMultipleChoiceExercise
key: 5fbf47b162
xp: 50
```

Was bewirkt der Ctrl-L im Konsolenfenster?

`@hint`
L wie Leer

`@possible_answers`
- Listet die zuletzt verwendeten Befehle auf
- [Löscht den Inhalt des Konsolenfensters]
- Laden von Dateien
- Löschen der Variablen im Environment

`@feedback`
- Sorry, ist falsch
- Sehr gut, nur weiter so!
- Sorry, ist falsch
- Sorry, ist falsch

---

## Pfeiltasten

```yaml
type: PureMultipleChoiceExercise
key: 6fe296d94b
xp: 50
```

Was bewirken die Pfeiltasten (Up/Down) im Konsolenfenster?

`@hint`
Erspart eventuell Tipparbeit.

`@possible_answers`
- Vergrößert das Konsolenfenster
- Verkleinert das Konsolenfenster
- [Zeigt die Historie der bereits verwendeten Befehle]
- Hat keine Wirkung

`@feedback`
- Sorry, ist falsch
- Sorry, ist falsch
- [Na also, passt ja alles.]
- Sorry, ist falsch

---

## Ausdrücke in R

```yaml
type: PureMultipleChoiceExercise
key: 89a188d0cb
xp: 50
```

Welche Bedeutung hat in R das folgende Symbol == im Ausdruck 8 * 3 == 6 * 4?

`@hint`


`@possible_answers`
- weist der Variablen _8 * 3_ das Ergebnis der Multiplikation _6 * 4_ zu.
- hat gar keine Bedeutung, da dieser Ausdruck in R gar nicht erlaubt ist.
- [Vergleicht, ob das Ergebnis 8 * 3 gleich 6 * 4 ist und liefert den logischen Wert TRUE als Ergebnis]
- Vergleicht, ob das Ergebnis 8 * 3 gleich 6 * 4 ist und liefert den logischen Wert FALSE als Ergebnis

`@feedback`
- Sorry, total daneben. Vielleicht solltest du das Skript nochmals von vorne zu lesen beginnen.
- Nein, der Ausdruck ist sehr wohl erlaubt.
- Ja, er prüft auf logische Äquivalenz.
- Der Vergleich liefert aber ein anderes Ergebnis, oder?

---

## Tab

```yaml
type: PureMultipleChoiceExercise
key: 0b6e4acaf0
xp: 50
```

Der Tabulator hat im Konsolenfenster eine besondere Eigenschaft. Welche der nachfolgenden Eigenschaften ist richtig?

`@hint`
Ergänzung von Befehlen

`@possible_answers`
- Fügt einen fest vorgegebenen Abstand an Leerzeichen ein.
- [Sucht nach einer Funktion, die mit denselben Buchstaben beginnt wie bereits eingegeben und zeigt diese Funktionen einerseits an, und vervollständigt den Namen des begonnenen Funktionsnamens]
-  Wechselt vom Konsolefenster zum Environmentfenster.
-  Wechselt vom Konsolefenster zum Filefenster.

`@feedback`
- Sorry, ist falsch
- Genau, das habe ich erwartet!
- Sorry, ist falsch
- Sorry, ist falsch

---

## Tab-completion

```yaml
type: MultipleChoiceExercise
key: 710bb8bb0e
xp: 50
```

Berechne den Mittelwert des Vektors **vector1** (der Vektor wurde bereits in den Arbeitsspeicher geladen und ist mit eben diesem Namen verfügbar). Wenn du in der Konsole _v1_ und _Eingabe_ eingibst, siehst du die Werte, die in diesem Vektor gespeichtert sind.

Verwende die Tab-completion von RStudio sowohl beim Funktionsnamen wie auch innerhalb der Funktion.

`@possible_answers`
- 5.4
- [7.6]
- 10.45
- 12.23

`@hint`
Verwende die Funktion mean() und die Tab-completion-Funktion von R-Studio

`@pre_exercise_code`
```{r}
vector1 <- c(3, 12, 13, 4, 6)
```

`@sct`
```{r}
msg1 <- "Nein, nicht richtig"
msg2 <- "Nice one!"
msg3 <- "Fast, aber doch nicht ganz"
msg4 <- "Ohje, daneben"
ex() %>% check_mc(2, feedback_msgs = c(msg1, msg2, msg3, msg4))
```

---

## Editor

```yaml
type: PureMultipleChoiceExercise
key: 0b33dcf8ec
xp: 50
```

Welche Vorteile hat man, wenn der Editor (Source pane) anstelle der Konsole für die Erstellung von R-Programmen verwendet wird?

`@hint`
Mehrere Befehle können in der Konsole nur nacheinander eingegeben und ausgeführt werden.

`@possible_answers`
- [Der Editor ermöglicht Speicherung und Ausführung mehrerer Befehlszeilen, sowie die Dokumentation eines R-Programms.]
- Ergebnisse eines Programms können neben dem Code mitgespeichert werden.
- Mit dem Editor können Graphiken besser dargestellt werden.
- Die Hilfefunktion für R-Funktionen ist nur im Editor verfügbar.
- Die Historie von bereits ausgeführten Befehlssequenzen steht nur für Funktionen, die im Editor verwendet wurden, zur Verfügung.

`@feedback`
- Ja
- Nein
- Nein
- Nein
- Nein

---

## Console als Rechner

```yaml
type: NormalExercise
key: 1d6545be34
xp: 100
```

Bei dieser Aufgabe geht es darum, die Konsole von R-Studio als einfachen Taschenrechner zu verwenden.

`@instructions`
Berechnen Sie die Addition der Zahlen 2+28

`@hint`
Ergebnis sollte 30 sein, wenn du was anderes heraus bekommst, dann passt was nicht

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
2+28
```

`@sct`
```{r}
ex() %>% check_code(c("2+28", "28+2"), fixed = TRUE)
```

---

## Konsolenfenster

```yaml
type: PureMultipleChoiceExercise
key: 80a18784d0
xp: 50
```

Wofür wird in R-Studio das Konsolenfenster verwendet?

`@hint`
Sieh mal im Web nach.

`@possible_answers`
- [Um direkte R-Befehle abzusenden.]
- Um die R-Dateien zu verwalten.
- Um den Ablauf von R-Programmen zu kontrollieren.
- Um R-Programme zu schreiben und diese in einem Programmfile abzuspeichern.

`@feedback`
- Super!
- Das kann man besser im Dateimanager machen
- Dazu verwendet man den Debugger
- Dazu verwendet man das "Source-Fenster"
