---
title: RStudio
description: 'Allgemeine Fragen zu RStudio'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## F체r was verwendet man R-Studio?

```yaml
type: PureMultipleChoiceExercise
key: 85cda17d46
xp: 50
```

In welchen der folgend genannten F채lle ist RStudio eine ideale Anwendung?

`@hint`
Vielleicht solltest du ganz kurz mal die Homepage von R-Studio ansehen:-)

`@possible_answers`
1. Um R-Programme zu schreiben, den Programmablauf zu testen und Projektbasiert zu arbeiten.

2. Um die Ausgabe von R-Programmen dazustellen.

3. Um Graphiken und Tabellen von R-Programmen APA-konform darzustellen.

4. Um R-Programme in lauff채hige .exe Dateien umzuwandeln.

`@feedback`
Nicht schlecht f체r den Anfang!
