---
title: Homework
description: 'erste Hausübung'
---

## Vektoren erzeugen

```yaml
type: NormalExercise
key: f6db413b93
xp: 100
```

<!-- Guidelines for contexts: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Erzeugen Sie einen Vektor x, der aus den natürlichen Zahlen von 1 bis 10 besteht
  Hinweis: c(1,2,3,4,5,6,7,8,9,10) gilt nicht, es gibt eine einfachere Lösung!

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- c(x:y)

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
x <- c(1:10)
```

`@sct`
```{r}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
