---
title: RStudio
description: 'Allgemeine Fragen zu RStudio'
---

## Anwendungsbereich R-Studio

```yaml
type: PureMultipleChoiceExercise
key: 85cda17d46
xp: 50
```

In welchen der folgend genannten Fälle ist RStudio eine ideale Anwendung?

`@hint`
Vielleicht solltest du ganz kurz mal die Homepage von R-Studio ansehen

`@possible_answers`
[Um R-Programme zu schreiben, den Programmablauf zu testen und Projektbasiert zu arbeiten.]

Um die Ausgabe von R-Programmen tabellarisch und graphisch dazustellen.

Um Graphiken und Tabellen von R-Programmen APA-konform darzustellen.

Um R-Programme in lauffähige .exe Dateien umzuwandeln.

`@feedback`
Nicht schlecht für den Anfang!

---

## Konsolenfenster

```yaml
type: PureMultipleChoiceExercise
key: 80a18784d0
xp: 50
```

Für was wird in R-Studio das Konsolenfenster verwendet?

`@hint`
Sieh mal im Web nach

`@possible_answers`
[Um direkte R-Befehle abzusenden.]
Um die R-Dateien zu verwalten.
Um den Ablauf von R-Programmen zu kontrollieren.
Um R-Programme zu schreiben und diese in einem Programmfile abzuspeichern.

`@feedback`
Super!

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: adfd848eb6
xp: 50
```

Warum funktioniert das mit den Multiple Choice nicht?

`@possible_answers`
- zu wenig Erfahrung
- zu blöd
- zu alt
- [Hilfe nicht gelesen]

`@hint`
Letzteres ist immer gut

`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
